```python
# Initialize Otter
import otter
grader = otter.Notebook("lab00.ipynb")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

<img src="cs88-logo.png" alt="CS C88C Logo" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Lab 0: Getting Started

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# 1) What Would Python Do? (WWPD)

<hr style="border: 1px solid #fdb515;" />

One component of lab assignments is to predict how the Python will behave. For the following questions, assign the variable to the **output** of various statements/expressions.

---

## 1a)

```python
>>> x = 20
>>> x + 2
```

```python
one_a = ...
```

```python
grader.check("q1a")
```

---

## 1b)

```python
>>> x
```

```python
one_b = ...
```

```python
grader.check("q1b")
```

---

## 1c) 

```python
>>> y = 5
>>> y = y + 3
>>> y * 2
```

```python
one_c = ...
```

```python
grader.check("q1c")
```

---

## 1d)

```python
>>> y + x
```

```python
one_d = ...
```

```python
grader.check("q1d")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# 2) Implementing the 2025 Function

<hr style="border: 1px solid #fdb515;" />

Labs will often ask you to implement functions. You should see a function called `twenty_twenty_five` below that has a blank
`return` statement. That blank is the only part you should change. Come up with the most creative expression that evaluates to 2025 using only numbers and the +, *, and - operators (or ** and % if you'd like)

```python
def twenty_twenty_five():
    return ______

twenty_twenty_five() # This line should output 2025 after being run
```

```python
grader.check("q2")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1a (points: 1)
```
>>> import hashlib
>>> hashlib.sha256(str(one_a).encode()).hexdigest()
'785f3ec7eb32f30b90cd0fcf3657d388b5ff4297f2f9716ff66e9b69c05ddd09'
```


Tests for q1b (points: 1)
```
>>> import hashlib
>>> hashlib.sha256(str(one_b).encode()).hexdigest()
'f5ca38f748a1d6eaf726b8a42fb575c3c71f1864a8143301782de13da2d9202b'
```


Tests for q1c (points: 1)
```
>>> import hashlib
>>> hashlib.sha256(str(one_c).encode()).hexdigest()
'b17ef6d19c7a5b1ee83b907c595526dcb1eb06db8227d650d5dda0a9f4ce8cd9'
```


Tests for q1d (points: 1)
```
>>> import hashlib
>>> hashlib.sha256(str(one_d).encode()).hexdigest()
'59e19706d51d39f66711c2653cd7eb1291c94d9b55eb14bda74ce4dc636d015a'
```


Tests for q2 (points: 1)
```
>>> twenty_twenty_five()
2025
```
