```python
# Initialize Otter
import otter
grader = otter.Notebook("disc10.ipynb")
```

<img src="dsm-logo.png" alt="Data Science Mentors Logo" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Discussion Week 10 – Control and Iteration

## Data 6, Fall 2024

In this week's discussion we will be going over control and iteration. These new topics will greatly improve how we define functions! Below the reference sheet contains a section about booleans, and another resource is python tutor which can visualize most functions step by step process.

Resources: 
* [Python Reference Sheet](https://data6.org/fa24/reference/) 
* [Python Tutor](https://pythontutor.com/python-compiler.html#mode=edit)

```python
# Imports
from datascience import *
import numpy as np
Table.interactive_plots()
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 1: Trick or Treat 🍬🎃

The `treats` array holds 100 randomly chosen items (both treats and non-treats) from the `items` list. Although Data 6 hasn’t yet introduced `np.random.choice`, imagine it like drawing slips of paper from a hat. Each slip has an item written on it, and we record the item each time we draw. We repeated this "drawing" 100 times to create our `treats` array!

```python
items = np.array(["chocolate", "gummy bear", "candy corn", "lollipop", "caramel", "coal"])
treats = np.random.choice(items, 100)
treats
```

<hr style="border: 1px solid #fdb515;" />

## Question 1a: Coal?

We notice that the `items` array has coal, which is most definitely not a treat. We want to count the number of coal in the `treats` array! Create the function `coal_counter`, which takes in one arguement, `treats_array`. It should return the number of coal in `treats`.


```python
def coal_counter(treats_array):
    counter = 0 
    ...
        ...
            ...
    ...

coal_counter(treats) # This line counts the number of coal in your treats array
```

```python
grader.check("q1a")
```

<hr style="border: 1px solid #fdb515;" />

## Question 1b: Chocolate Lover 🍫

Edwin is a chocolate enjoyer. He wants to go through his treats, but he some interesting reactions. He screams "CHOCOLATE" everytime he picks up a chocalate, and "Meh" for everything else. Write a function, `edwins_reactions` which takes in an array, `treats_array`, and prints his reaction for every item in `edwins_treats`.

*Note*: This function doesn't return anything

```python
edwins_treats = np.random.choice(items, 10) # Edwin only has 10 pieces of candy

def edwins_reactions(treats_array):
    for ...:
        ...
            ...
        ...
            ...

edwins_reactions(edwins_treats) # Edwin's reactions to his treats
```

```python
grader.check("q1b")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 2: San Francisco Salaries 2023

Let's take a look at the salaries of employees of the City of San Francisco in 2023. The median salary reported by the city government was about $169,361. **Note that in this lab, this dataset is the population.**

This data was collected by [Transparent California](https://transparentcalifornia.com/salaries/2023/san-francisco/).

```python
salaries = Table.read_table("sf_salaries_2023.csv")
salaries
```

<!-- BEGIN QUESTION -->

<hr style="border: 1px solid #fdb515;" />

## Question 2a: Exploring the data

How is `Total Pay` calculated?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<hr style="border: 1px solid #fdb515;" />

## Question 2b: Overlaid Visualizations

Create and overlaid visualization of your choosing. In the solutions, I picked the two most popular job titles and made an overlaid histogram on their total pay. 

```python
...
```

<!-- END QUESTION -->

## We are almost 2/3 of the ways done with the semester!

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1a (points: 1)
```
>>> np.random.seed(42)
>>> treats = np.random.choice(items, 100)
>>> coal_counter(treats) == 17
True
```


Tests for q1b (points: 1)
```
>>> callable(edwins_reactions)
True
```
