```python
# Initialize Otter
import otter
grader = otter.Notebook("disc08.ipynb")
```

<img src="dsm-logo.png" alt="Data Science Mentors Logo" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Discussion Week 8 – Groups and Visualizations

## Data 6, Fall 2024

Pivots and groups have plenty of resources for visualizing how both of these work. One small disclaimer is that, the **table function visualizer** is a bit innaccurate. It forgets to sort the table in ascending order. Hope these help though!

Data 6 goes more into depth on visualizations compared to Data 8. The questions on this notebook are some examples, but will not cover every scenario.

Resources: 
* [Python Reference Sheet](https://data6.org/fa24/reference/) 
* [Table Function Visualizer](https://www.data8.org/interactive_table_functions/) for visualizing `.group` and `.pivot`
* [Group Animation](https://www.youtube.com/watch?v=HLoYTCUP0fc&ab_channel=Data8)
* [Pivot Animation](https://www.youtube.com/watch?v=4WzXo8eKLAg&ab_channel=Data8)

```python
# Imports
from datascience import *
import numpy as np
Table.interactive_plots()
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 1: YouTube Eats!

---

The `ratings` table contains 150 ratings on popular Youtuber food products. 

* `Item`: a _string_, the name of the food product 
* `Gender`: a _string_, the person's Gender
* `Age`: a _int_, the age of the person who is reviewing the item
* `Rating`: a _float_, the rating of the item (1-10)

```python
ratings = Table.read_table("ratings.csv")
ratings
```

<hr style="border: 1px solid #fdb515;" />

## Question 1a: Average Overall Ratings

What is the combined average rating for the `ratings` table?


```python
combined_average = ...
combined_average
```

```python
grader.check("q1a")
```

<hr style="border: 1px solid #fdb515;" />

## Question 1b: Average Rating for Each Item!

Now that we have seen the average rating for all groups combined, we're now interested in the average rating for each `Item`. Assign `item_averages` to an array of the averages of items. 

*Hint*: `tbl.group()` may be helpful.

```python
item_averages = ...
item_averages
```

```python
grader.check("q1b")
```

<!-- BEGIN QUESTION -->

<hr style="border: 1px solid #fdb515;" />

## Question 1c: Visualizations

We wish to visualize the **correlation** between `Age` and `ratings`. What would be the most appropriate visualization? 

*Note:* Your options are line plot, scatter plot, histogram, and bar charts. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<hr style="border: 1px solid #fdb515;" />

## Question 1d: Now visualize!

Create two visualizations! One that accomplishes the goal from above. That is a visualization with `Age` on the x-axis and `Rating` on the y-axis. Then another visualization that encodes `Item`.

```python
...
...
```

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 2: NBA

Here, the `nba` table contains player data from all players in the nba in 2023-2024. This is the same data from discussion 7.

```python
nba = Table.read_table("nba.csv")
nba
```

<!-- BEGIN QUESTION -->

<hr style="border: 1px solid #fdb515;" />

## Question 2a: Bar Charts

Recall that Barcharts uses one categorical data and one numerical data point. Ideally, the groups should be aggregated (average, sum, count, etc). Now, create a bar chart on the count of teams.

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<hr style="border: 1px solid #fdb515;" />

## Question 2b: Sorting our bar charts

Do the same barh chart, but now sort the barh chart in descending order of `count`.

```python
...
```

<!-- END QUESTION -->

## That's all, congrats on all the hard work!

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1a (points: 1)
```
>>> np.isclose(combined_average, 4.02406)
True
```


Tests for q1b (points: 1)
```
>>> isinstance(item_averages, np.ndarray)
True
```
```
>>> sum(item_averages) == 28.162423071046604
True
```
